<a href="https://colab.research.google.com/github/Juma1611/Humphrey-/blob/main/Copy_of_Assignment_Control_Statements.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 02.1 (if-else)

Consider the following assessment criteria which map a score out of 100 to an
assessment grade:

| Grade            | Raw score  (/100)      |
| ---------------- | ---------------------- |
| Excellent        | $\ge 85$               |
| Very good        | $\ge 76.5$ and $< 85$  |
| Good             | $\ge 64$ and $< 76.5$  |
| Need improvement | $\ge 40$ and $< 64$    |
| Did you try?     | $< 40$                 |

Write a program that, given an a score, prints the appropriate grade. Print an error message if the input score is greater than 100 or less than zero.

```python
def grade(score):
    if score < 0 or score > 100:
        return "Invalid score"
    elif score >= 85:
        return "Excellent"
    elif score >= 76.5:
        return "Very good"
    elif score >= 64:
        return "Good"
    elif score >= 40:
        return "Need improvement"
    else:
        return "Did you try?"

print(grade(72))  # Output: Need improvement
```

Output:
```
Good
```

## Exercise 02.2 (bisection)

Bisection is an iterative method for finding approximate roots of a function. Say we know that the function $f(x)$ has one root between $x_{0}$ and $x_{1}$ ($x_{0} < x_{1}$). We then:

- Evaluate $f$ at the midpoint $x_{\rm mid} = (x_0 + x_1)/2$, i.e. compute
   $f_{\rm mid} = f(x_{\rm mid})$
- Evaluate $f(x_0) \cdot f(x_{\rm mid})$

  - if $f(x_0) \cdot f(x_{\rm mid}) < 0$:

    $f$ must change sign somewhere between $x_0$ and $x_{\rm mid}$, hence the root must lie between
    $x_0$ and $x_{\rm mid}$, so set $x_1 = x_{\rm mid}$.
   
  - else:

    $f$ must change sign somewhere between $x_{\rm mid}$ and $x_1$, so set
    $x_0 = x_{\rm mid}$.

The above steps can be repeated a specified number of times, or until $|f_{\rm mid}|$
is below a tolerance, with $x_{\rm mid}$ being the approximate root.


### Task

The function

$$
f(x) = \frac{x^{5}}{10} + x^3 - 10x^2 + 4x + 7
$$


has one root in the range $0 < x < 2$.

1. Use the bisection method to find an approximate root $x_{r}$ using 20 iterations
   (use a `for` loop).
2. Use the bisection method to find an approximate root $x_{r}$ such that
   $\left| f(x_{r}) \right| < 1 \times 10^{-6}$ and report the number of iterations
   required (use a `while` loop).

Store the approximate root using the variable `x_mid`, and store $f(x_{\rm mid})$ using the variable `f`.

*Hint:* Use  `abs` to compute the absolute value of a number, e.g. `y = abs(x)` assigns the absolute value of `x` to `y`.

#### (1) Using a `for` loop.

<!-- grade_id: cell-3a0d0a59613d6b3f, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

# Use 20 iterations
for n in range(20):
    # Compute midpoint
    x_mid = (x0 + x1) / 2

    # Evaluate function at (i) left end-point and at (ii) midpoint
    f0 = (x0**5) / 10 + x0**3 - 10 * x0**2 + 4 * x0 + 7
    f = (x_mid**5) / 10 + x_mid**3 - 10 * x_mid**2 + 4 * x_mid + 7

    ...

    print(n, x_mid, f)
```

Output:
```
0 1.0 2.0999999999999996
1 1.0 2.0999999999999996
2 1.0 2.0999999999999996
3 1.0 2.0999999999999996
4 1.0 2.0999999999999996
5 1.0 2.0999999999999996
6 1.0 2.0999999999999996
7 1.0 2.0999999999999996
8 1.0 2.0999999999999996
9 1.0 2.0999999999999996
10 1.0 2.0999999999999996
11 1.0 2.0999999999999996
12 1.0 2.0999999999999996
13 1.0 2.0999999999999996
14 1.0 2.0999999999999996
15 1.0 2.0999999999999996
16 1.0 2.0999999999999996
17 1.0 2.0999999999999996
18 1.0 2.0999999999999996
19 1.0 2.0999999999999996
```

#### (2) Using a `while` loop

Use the variable `counter` for the iteration number.

*Remember to guard against infinite loops.*

<!-- grade_id: cell-1747a3fb1dd1de48, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

tol = 1.0e-6
error = tol + 1.0

# Iterate until tolerance is met
counter = 0
while error > tol:
    ...

    # Guard against an infinite loop
    if counter > 1000:
        print("Oops, iteration count is very large. Breaking out of while loop.")
        break

    print(counter, x_mid, error)
```

<!-- grade_id: cell-ddff32f44df88388, points: 0, grade: true, solution: false -->
```python
## tests ##
assert counter == 24
assert abs(f) < 1.0e-6
```

## Exercise 02.3 (series expansion)

For $|x| < 1$ the series:

$$
(1 + x)^{-1/2} = \sum_{n = 0}^{\infty} \frac{(-1)^n (2n)!}{4^n (n!)^2} x^n
$$

converges.

1. Using a `for` statement, approximate $1/\sqrt{0.16}$ using 30 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $1/\sqrt{0.16}$ to within $1 \times 10^{-5}$.

Store the absolute value of the error in the variable `error`.

### Hints

To compute the factorial, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
```
You only need `import math` once at the top of your program. Standard modules, like `math`, will be explained in a later

<!-- The power series expansion for the sine function is:

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

(See mathematics data book for a less compact version; this compact version is preferred here as it is simpler to program.)

1. Using a `for` statement, approximate $\sin(3\pi/2)$ using 15 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $\sin(3\pi/2)$ to within $1 \times 10^{-8}$.

Store the absolute value of the error in the variable `error`.

*Note:* Calculators and computers use iterative or series expansions to compute trigonometric functions, similar to the one above (although they use more efficient formulations than the above series).

### Hints

To compute the factorial and to get a good approximation of $\pi$, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
pi = math.pi
```
You only need '`import math`' once at the top of your program. Standard modules, like `math`, will be explained in a later. If you want to test for angles for which sine is not simple, you can use
```python
a = 1.3
s = math.sin(a)
```    
to get an accurate computation of sine to check the error. -->

```python
def f(x):
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7

x0, x1 = 0.0, 2.0
for n in range(20):
    x_mid = (x0 + x1) / 2
    f_mid = f(x_mid)
    if f(x0) * f_mid < 0:
        x1 = x_mid
    else:
        x0 = x_mid
    print(f"Iteration {n+1}: x_mid = {x_mid}, f(x_mid) = {f_mid}")
```

Output:
```
Iteration 1: x_mid = 1.0, f(x_mid) = 2.0999999999999996
Iteration 2: x_mid = 1.5, f(x_mid) = -5.365625000000001
Iteration 3: x_mid = 1.25, f(x_mid) = -1.36669921875
Iteration 4: x_mid = 1.125, f(x_mid) = 0.4477813720703132
Iteration 5: x_mid = 1.1875, f(x_mid) = -0.4408627510070797
Iteration 6: x_mid = 1.15625, f(x_mid) = 0.008327355980872753
Iteration 7: x_mid = 1.171875, f(x_mid) = -0.21507753478363156
Iteration 8: x_mid = 1.1640625, f(x_mid) = -0.1030741602036862
Iteration 9: x_mid = 1.16015625, f(x_mid) = -0.04729774910292761
Iteration 10: x_mid = 1.158203125, f(x_mid) = -0.019466230897839054
Iteration 11: x_mid = 1.1572265625, f(x_mid) = -0.005564689502373099
Iteration 12: x_mid = 1.15673828125, f(x_mid) = 0.0013825210451265946
Iteration 13: x_mid = 1.156982421875, f(x_mid) = -0.0020907873792292975
Iteration 14: x_mid = 1.1568603515625, f(x_mid) = -0.00035405894194262544
Iteration 15: x_mid = 1.15679931640625, f(x_mid) = 0.0005142496094645566
Iteration 16: x_mid = 1.15682983398437
```

#### (1) Using a `for` loop

<!-- grade_id: cell-d6df4b428f969a5d, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.factorial
import math

# Value of x (such that (1 + x) = 0.16
x = -0.84

# Initialise approximation of the function
approx_f = 0.0

...

print("The error is:")
print(error)
```

<!-- grade_id: cell-8b269b717e92414d, points: 0, grade: true, solution: false -->
```python
def f(x):
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7

x0, x1 = 0.0, 2.0
for n in range(20):
    x_mid = (x0 + x1) / 2
    f_mid = f(x_mid)
    if f(x0) * f_mid < 0:
        x1 = x_mid
    else:
        x0 = x_mid
    print(f"Iteration {n+1}: x_mid = {x_mid}, f(x_mid) = {f_mid}")
```

Output:
```
Iteration 1: x_mid = 1.0, f(x_mid) = 2.0999999999999996
Iteration 2: x_mid = 1.5, f(x_mid) = -5.365625000000001
Iteration 3: x_mid = 1.25, f(x_mid) = -1.36669921875
Iteration 4: x_mid = 1.125, f(x_mid) = 0.4477813720703132
Iteration 5: x_mid = 1.1875, f(x_mid) = -0.4408627510070797
Iteration 6: x_mid = 1.15625, f(x_mid) = 0.008327355980872753
Iteration 7: x_mid = 1.171875, f(x_mid) = -0.21507753478363156
Iteration 8: x_mid = 1.1640625, f(x_mid) = -0.1030741602036862
Iteration 9: x_mid = 1.16015625, f(x_mid) = -0.04729774910292761
Iteration 10: x_mid = 1.158203125, f(x_mid) = -0.019466230897839054
Iteration 11: x_mid = 1.1572265625, f(x_mid) = -0.005564689502373099
Iteration 12: x_mid = 1.15673828125, f(x_mid) = 0.0013825210451265946
Iteration 13: x_mid = 1.156982421875, f(x_mid) = -0.0020907873792292975
Iteration 14: x_mid = 1.1568603515625, f(x_mid) = -0.00035405894194262544
Iteration 15: x_mid = 1.15679931640625, f(x_mid) = 0.0005142496094645566
Iteration 16: x_mid = 1.15682983398437
```

#### (2) Using a `while` loop

*Remember to guard against infinite loops.*

<!-- grade_id: cell-03df25d83df3674e, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.sin and math.factorial
import math

# Value of x (such that (1 - x) = 0.16)
x = -0.84

# Tolerance and initial error (this just needs to be larger than tol)
tol = 1.0e-5
error = tol + 1.0

# Initialise approximation of function
approx_f = 0.0

# Initialise counter
n = 0

# Loop until error satisfies tolerance, with a check to avoid
# an infinite loop
while error > tol and n < 1000:

    ...

    # Increment counter
    n += 1


print("\nThe error is:", error)
print("Number of terms in series:", n)
```

<!-- grade_id: cell-03df25d83df3674dfwe, points: 0, grade: true, solution: false -->
```python
x0, x1 = 0.0, 2.0
tol = 1e-6
counter = 0
while True:
    x_mid = (x0 + x1) / 2
    f_mid = f(x_mid)
    if abs(f_mid) < tol or counter >= 1000:
        break
    if f(x0) * f_mid < 0:
        x1 = x_mid
    else:
        x0 = x_mid
    counter += 1
print(f"Root: x = {x_mid}, f(x) = {f_mid}, iterations = {counter}")
```

Output:
```
Root: x = 1.1568354368209839, f(x) = 3.9180508970559913e-07, iterations = 23
```